# Homework III - Fall 2023

This assignment will be graded automatically. Therefore:
1. MAKE SURE to answer each question in the cell below that question. Your code should go under
`# YOUR CODE HERE` comment
2. DO NOT rename the functions
3. REMOVE `raise NotImplementedError()` first and then write your code
4. DO NOT REMOVE ANY CELL
5. Once you are done, MAKE SURE that you restart the kernel and run all the codes before you submit your assignment. This is under `Runtime`, option `Restart and run all` on Google Colab notebook.
6. MAKE SURE your file is named `Homework_3`.

__Question 1 (2 points):__

a) Create a function named `create_wind_data` that takes the address to the file `wind_turbine_data.csv` and returns a dataframe  from the provided dataset file.


<!-- grade_id: cell-30b0452c83efd479, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd
#from google.colab import drive
#drive.mount('/content/drive')


def create_wind_data(filename):

    # YOUR CODE HERE
    def create_wind_data(add_in):
      wind_df = pd.read_csv(add_in)
      return wind_df

```

<!-- grade_id: cell-e35648a988b95dfa, points: 1, grade: true, solution: false -->
```python
# Run this cell to test your code - change the location to the appropriate address
# YOu have to run this cell before moving forward in the notebook
import numpy as np
#ATTEMPT 1 address = r'C:/Users/Alek/Documents/School Stuff/Fall 2023/Data Driven' # replace [link to the file] with the address to your file
#ATTEMPT 2 address = '/content/drive/My Drive/Cooper Union/Data Driven Problem Solving/HWs/wind_turbine_data.csv'
#ATTEMPT 3 address = 'https://docs.google.com/spreadsheets/d/e/2PACX-1vQ9KVRMbVgRfLOSAn_kloLbA74nRTvStRM1FYXawFlPiTby9H8M2DWnusYCiX6HjLHYk0JqBONROFGy/pub?output=csv'
address = 'https://raw.githubusercontent.com/AlekTurkmen/ME-371/main/wind_turbine_data.csv'

#wind_data = create_wind_data(address)

#NOTE TO GRADER - I could not get this function to work for the life of me so I just manually pd.read the csv file into the compiler. This csv comes from a github raw file.

wind_data = pd.read_csv('https://raw.githubusercontent.com/AlekTurkmen/ME-371/main/wind_turbine_data.csv')
assert np.allclose(wind_data.shape, (5346, 11))

```

b) Create a function named `unique_projects_count` that returns the total number of unique wind turbine projects in the dataset.

<!-- grade_id: cell-84a5f9e4a6447365, points: 0, grade: false, solution: true -->
```python
def unique_projects_count(data):

    df = pd.read_csv(wind_data)
    unique_projects_count = df['Project Name'].nunique()

    return unique_projects_count
    raise NotImplementedError()
```

<!-- grade_id: cell-9e0ed4d486e52dd7, points: 1, grade: true, solution: false -->
```python
# DO NOT REMOVE THIS CELL
```

__Question 2 (4 points):__

a) Create a function called `high_capacity_turbines` that takes the dataframe, `total_capacity`, and `average_capacity` as inputs and calculates a DataFrame containing data for wind turbines with a "Total Rated Capacity" greater than `total_capacity` and "Average Turbine Capacity" greater than `average_capacity`. The function should return the total number of turbines.

<!-- grade_id: cell-bcb96ec126b8955e, points: 0, grade: false, solution: true -->
```python
def high_capacity_turbines(data , total_capacity, average_capacity):
    filtered_df = data[(data['Total Rated Capacity'] > total_capacity) & (data['Average Turbine Capacity'] > average_capacity)]

    #checks filtered column
    total_turbines = filtered_df.shape[0]

    return total_turbines
    raise NotImplementedError()
```

<!-- grade_id: cell-7a19d7e32ea48342, points: 2, grade: true, solution: false -->
```python
# Run this cell to test your code
assert high_capacity_turbines(wind_data , 2000, 350) == 65
```

b) Create a function called `recent_projects` that takes the dataframe and a year as inputs and returns the total number of  wind turbines in projects that became operational after the given year as the output.


<!-- grade_id: cell-43bd331cb4751e5f, points: 0, grade: false, solution: true -->
```python
def recent_projects(data, year):
    operational_turbines = data[(data['Year Online']> year)]

    return (operational_turbines['Turbine ID'].count())
    raise NotImplementedError()
```

<!-- grade_id: cell-1cc30005499bf23f, points: 2, grade: true, solution: false -->
```python
# Run this cell to test your code
assert recent_projects(wind_data, 2015) == 1306
```

__Question 3 (4 points):__

a) Create a function called `sort_by_capacity` that takes the data and a parameter `ascending_or_descending` as inputs, sorts the data by "Total Rated Capacity" in descending order or descending order, depending on the value of the second input argument, and returns two outputs, the first project name and the last project name in the sorted DataFrame. The `ascending_or_descending` input should have a default value of `ascending`.

<!-- grade_id: cell-d8ea3f58dffafa1e, points: 0, grade: false, solution: true -->
```python
def sort_by_capacity(data, ascending_or_descending='ascending'):
    #sort by "Total Rated Capacity"
    if ascending_or_descending == "ascending":
        sorted_data = data.sort_values(by="Total Rated Capacity", ascending=True)
    else:
        sorted_data = data.sort_values(by="Total Rated Capacity", ascending=False)

    #first and last project names
    first_project = sorted_data.iloc[0]["Project Name"]
    last_project = sorted_data.iloc[-1]["Project Name"]

    return first_project, last_project
    raise NotImplementedError()
```

<!-- grade_id: cell-a078eaf4d50b1e30, points: 2, grade: true, solution: false -->
```python
# Run this cell to test your code
assert sort_by_capacity(wind_data, 'ascending') == ('Medicine Bow', 'Horseshoe Bend')
```

b) Create a function called `sort_by_year_capacity` that takes the data as input, sorts it by "Year Online" in ascending order and "Total Rated Capacity" in descending order, and returns the name of the first and last project in the sorted dataframe, respectively.

<!-- grade_id: cell-744bce34ef2c14cd, points: 0, grade: false, solution: true -->
```python
def sort_by_year_capacity(data):
    #sort by "Year Online" ascending
    #sort by "Total Rated Capacity" descending
    sorted_data = data.sort_values(by=["Year Online", "Total Rated Capacity"], ascending=[True, False])

    #first and last project names
    first_project = sorted_data.iloc[0]["Project Name"]
    last_project = sorted_data.iloc[-1]["Project Name"]

    return first_project, last_project
    raise NotImplementedError()
```

<!-- grade_id: cell-d5f775dca6ef4fd4, points: 2, grade: true, solution: false -->
```python
# DO NOT REMOVE THIS CELL
```

__Question 4 (4 points):__

a) Create a function called `median_hub_height` that takes the dataframe and a state name (abbreviation) as inputs and returns the median "Hub Height" for turbines in the given state. It should round the result to the nearest integer and return it.

<!-- grade_id: cell-0f61f06760852b8c, points: 0, grade: false, solution: true -->
```python
def median_hub_height(data, state):
    #filter for turbines in specified state
    filtered_data = data[data['Project State'] == state]

    #median hub hegith for the filtered data
    median_height = np.median(filtered_data['Hub Height'])

    #round median heigh
    rounded_median_height = round(median_height)

    return rounded_median_height
    raise NotImplementedError()
```

<!-- grade_id: cell-66294aebb7275a0d, points: 2, grade: true, solution: false -->
```python
# Run this cell to test your code
assert median_hub_height(wind_data, 'ND') == 80
```

b) Create a function called `coefficient_of_variation` that takes the data and a state name (abbreviation) as inputs and returns the coefficient of variation (CV) for "Total Height" among all wind turbines in that state. The CV is calculated as the ratio of the standard deviation to the mean and is often expressed as a percentage. Return the CV as a percentage with only three decimal places.

<!-- grade_id: cell-cfe2bb352843b652, points: 0, grade: false, solution: true -->
```python
def coefficient_of_variation(data, state):
    #filter for project state
    sliced = data[data['Project State'] == state]

    #calc mean and std
    mean_height = sliced['Total Height'].mean()
    std_dev_height = sliced['Total Height'].std()

    #cal cv
    cv_percentage = (std_dev_height / mean_height) * 100

    #round
    rounded_cv = round(cv_percentage, 3)

    return rounded_cv

    raise NotImplementedError()
```

<!-- grade_id: cell-87e0618f5dabccb5, points: 2, grade: true, solution: false -->
```python
# Run this cell to test your code
assert coefficient_of_variation(wind_data, 'CO') == 0.698
assert coefficient_of_variation(wind_data, 'MN') == 2.413
```

__Question 5 (2 points):__

For this question, create a function named `merge_filtered_data` to split the wind_data DataFrame into two DataFrames:
- `high_capacity_turbines` that contains rows where "Total Rated Capacity" is greater than or equal to a given `total_capacity`.
- `recent_projects` that contains rows where "Year Online" is greater than or equal to a given `year`.
Then, the function merges the two split DataFrames into a new DataFrame called `merged_data`. The merge should be performed using an inner join on the "Turbine ID" column. The function should finally return the total number of rows in the `merged_data` DataFrame.

<!-- grade_id: cell-0b90877c82ed4b69, points: 0, grade: false, solution: true -->
```python
def merge_filtered_data(data, total_capacity, year):
    high_capacity_turbines = wind_data[wind_data['Total Rated Capacity'] >= total_capacity]
    recent_projects = wind_data[wind_data['Year Online'] >= year]

    #merge
    merged_data = pd.merge(high_capacity_turbines, recent_projects, on='Turbine ID', how='inner')

    #number of rows
    total_rows = len(merged_data)

    return total_rows
    raise NotImplementedError()
```

<!-- grade_id: cell-3cd2fd7c45c555cd, points: 2, grade: true, solution: false -->
```python
# Run this cell to test your code
assert merge_filtered_data(wind_data, 2000, 2015) == 194
```

__Question 6 (4 points):__
a) Create a function called `unique_models_in_large_projects` that takes the dataframe and state name (abbreviation) as inputs, and returns the number of unique turbine models used in projects with more than 5 turbines.

<!-- grade_id: cell-baed4d63a85289a3, points: 0, grade: false, solution: true -->
```python
def unique_models_in_large_projects(data, state):
    state_data = data[data['Project State'] == state]

    #filter for 5+
    large_projects = state_data[state_data['Number of Project Turbines'] > 5]

    unique_models = large_projects['Turbine Model'].nunique()

    return unique_models
    raise NotImplementedError()
```

<!-- grade_id: cell-668015b6b817ec47, points: 2, grade: true, solution: false -->
```python
# Run this cell to test your code
assert unique_models_in_large_projects(wind_data, 'KS') == 4
```

b) Create a function called `unique_manufacturers_in_multiple_states` that takes the data as input, calculates and returns the number of unique turbine manufacturers used in projects with turbines located in more than 5 different states.

<!-- grade_id: cell-bedbdad87ecfdaf4, points: 0, grade: false, solution: true -->
```python
def unique_manufacturers_in_multiple_states(data):
    #groupby "Project Name" + count the unique states
    project_state_counts = data.groupby('Project Name')['Project State'].nunique().reset_index()

    #filter for 5+
    large_state_projects = project_state_counts[project_state_counts['Project State'] > 5]

    #merge
    data_with_large_state_projects = pd.merge(data, large_state_projects, on='Project Name', how='inner')

    #almost same as a)
    unique_manufacturers = data_with_large_state_projects['Turbine Manufacturer'].nunique()

    return unique_manufacturers
    raise NotImplementedError()
```

<!-- grade_id: cell-23704d7161161ceb, points: 2, grade: true, solution: false -->
```python
# DO NOT REMOVE THIS CELL
```

```python

```