# Homework I - Fall 2023

This assignment will be graded automatically. Therefore:
1. MAKE SURE to answer each question in the cell below that question and your code goes under
`# YOUR CODE HERE` comment
2. DO NOT rename the functions
3. REMOVE `raise NotImplementedError()` first and then write your code
4. DO NOT REMOVE ANY CELL
5. Once you are done, MAKE SURE that you restart the kernel and run all the codes before you submit your assignment. This is under `Runtime`, option `Restart and run all` on Google Colab notebook.
6. DO NOT rename the file.

**Question I (2 pts):** Write a function that takes in a list and returns a list with items that are squared of items in the original list. Don't use any special Python library. For example,
```python
square_list([10])
```
should return
```python
[100]
```

<!-- grade_id: cell-198c9585a2609842, points: 0, grade: false, solution: true -->
```python
def square_list(input_list):
    """
    The input to the function will be a list.
    Calling the function on a list should return a list with
    squared values as its items.
    """
    # YOUR CODE HERE
    squared_list = []  #init empty list
    for item in input_list:
        squared_list.append(item ** 2)  #square items
    return squared_list
```

<!-- grade_id: cell-867298ad0eae73db, points: 2, grade: true, solution: false -->
```python
# Run this cell to test your code
assert square_list([10]) == [100]
```

**Quesion II (4 pts):**

_Part (1)_: Write a function that takes a dictionary and a key-value as inputs, add the key-value to the dictionary, and return the updated dictionary. For example, if the function take
```python
SampleDictionary = {'Data-Driven Prb Solving' : '9am', 'Pyhon for Data Analysis' : '6:10pm', 'Reinforcement Learning' : '10am'}
```
and
```python
'added key' : 'corresponding value'
```
it should return
```python
{'Data-Driven Prb Solving' : '9am', 'Pyhon for Data Analysis' : '6:10pm', 'Reinforcement Learning' : '10am', 'added key' : 'corresponding value'}
```

<!-- grade_id: cell-86b7bf9cb9455a70, points: 0, grade: false, solution: true -->
```python
def adding_key_values(dictionary_input, key, value):
    """
    The inputs are:
    dictionary_input: original dictionary with some keys and values
    key: the key that we want to add to the dictionary
    value: the corresponding value for the given key

    The output:
    the updated dictionary with added key-value pair
    """
    # YOUR CODE HERE
    dictionary_input.update({key: value})#adds key and value (2nd and 3rd inputs) into dictionary (1st input)
    return dictionary_input

```

<!-- grade_id: cell-f4f2476d0c713e1a, points: 2, grade: true, solution: false -->
```python
# Run this cell to test your code
assert adding_key_values({'key1':'value1'}, 'key2', 'value2')=={'key1':'value1', 'key2':'value2'}
```

_Part (2)_: Write a new function that takes a dictionary and a key, removes the key and its corresponding value from the dictionary, and then returns the updated dictionary. For example, if the function takes
```python
{'Data-Driven Prb Solving' : '9am', 'Pyhon for Data Analysis' : '6:10pm', 'Reinforcement Learning' : '10am'}
```
and
```python
'Data-Driven Prb Solving'
```
it should return
```python
{'Pyhon for Data Analysis' : '6:10pm', 'Reinforcement Learning' : '10am'}
```

<!-- grade_id: cell-5d7a969a6e3c6338, points: 0, grade: false, solution: true -->
```python
def removing_a_key_value(dictionary_input, key):
    """
    The inputs are:
    dictionary_input: original dictionary with some keys and values
    key: the key that we want to remove with its value

    The output:
    the updated dictionary without the key-value pair
    """
    # YOUR CODE HERE
    dictionary_input.pop(key)#deletes key (2nd input) out of dictionary (1st input)
    return dictionary_input

```

<!-- grade_id: cell-72609d9ab586bc9a, points: 2, grade: true, solution: false -->
```python
# Run this cell to test your code
assert removing_a_key_value(
    {'Data-Driven Prb Solving' : '9am',
     'Pyhon for Data Analysis' : '6:10pm',
     'Reinforcement Learning' : '10am'}, 'Data-Driven Prb Solving') == {
    'Pyhon for Data Analysis' : '6:10pm',
    'Reinforcement Learning' : '10am'}
```

**Question III (3 pts):** Write a function that takes a string, converts the string into a list, and then returns a list with items that are evey other letter in the string. For example, if the function takes
```python
'This string is written for homework I'
```
the output should be
```python
['T', 'i', ' ', 't', 'i', 'g', 'i', ' ', 'r', 't', 'e', ' ', 'o', ' ', 'o', 'e', 'o', 'k', 'I']
```

<!-- grade_id: cell-13216cb4a3cea148, points: 0, grade: false, solution: true -->
```python
def string_manipulation(string_in):
    """
    The input:
    string_in: a string

    The output:
    a list with every other letter of the given string as its items
    """
    # YOUR CODE HERE
    char_list = list(string_in) #convert to string to list
    string_list  = [] #init empty string
    for letter in range(0, len(char_list), 2):#iterate every 2
      string_list.append(char_list[letter]) #add letter into list

    return string_list
```

<!-- grade_id: cell-53a09fabefab8c11, points: 3, grade: true, solution: false -->
```python
# Run this cell to test your code
assert string_manipulation('This string is written for homework I') == ['T', 'i', ' ', 't', 'i', 'g', 'i', ' ', 'r', 't', 'e', ' ', 'o', ' ', 'o', 'e', 'o', 'k', 'I']
```

**Question IV (3):** Write a function that adds numbers from 1 to N, where N is the input to the function and the sum from 1 to N is the output.  For example if it takes 100, it should return 5050.

<!-- grade_id: cell-557401c87e177327, points: 0, grade: false, solution: true -->
```python
def sum_func(N):
    """
    The input:
    N is the number we want the summation from 1 to N

    The output:
    The summation of numbers from 1 to N
    """
    # YOUR CODE HERE
    total_sum = 0 #init counter
    for i in range(1, N + 1): #repeats N times
      total_sum += i #adds current iteration to counter

    return total_sum
```

<!-- grade_id: cell-b694531c530a04e0, points: 3, grade: true, solution: false -->
```python
# Run this cell to test your code
assert sum_func(100) == 5050
```

**Question V (3 pts):** The Chudnovsky series for approximating $\pi$ was used to obtain the world record for digits of $\pi$, 2.7 trillion digits:

$$
\frac{1}{\pi}=12 \sum_{k=0}^N\frac{(-1)^k(6k)!(13591409+545140134k)}{(3k)!(k!)^3(640320^{3k+3/2})}
$$

Write a function that takes N as an input and returns the estimated value of $\pi$ using this series.

<font color='blue'>Note:</font> The  function `math.factorial` can be used to calculate factorials in this fraction by calling it from library `math`
```python
import math
# factorial of number k
math.factorial(k)
```

<!-- grade_id: cell-280f31fc95308d55, points: 0, grade: false, solution: true -->
```python
def chudnovsky(N):
    """
    The input:
    N for the Chudnovsky series

    The output:
    estimated value of pi
    """
    # YOUR CODE HERE
    import math
    total_sum = 0.0 #init sigma sum
    k = 0 #init counter for N
    while k <= N:
        numerator = (-1) ** k * (math.factorial(6 * k) * (13591409 + 545140134 * k))
        denominator = math.factorial(3 * k) * (math.factorial(k) ** 3) * (640320 ** (3 * k + 3/2))
        total_sum = total_sum + (numerator / denominator)
        k = k + 1

    total_sum = 1 / (12 * total_sum) #reciprocal
    return total_sum
```

<!-- grade_id: cell-ed8b0c71a0bf563d, points: 3, grade: true, solution: false -->
```python
# Run this cell to test your code
assert chudnovsky(4) == 3.1415926535897936
```