# Homework II - Fall 2023

This assignment will be graded automatically. Therefore:
1. MAKE SURE to answer each question in the cell below that question. Your code should go under
`# YOUR CODE HERE` comment
2. DO NOT rename the functions
3. REMOVE `raise NotImplementedError()` first and then write your code
4. DO NOT REMOVE ANY CELL
5. Once you are done, MAKE SURE that you restart the kernel and run all the codes before you submit your assignment. This is under `Runtime`, option `Restart and run all` on Google Colab notebook.
6. DO NOT rename the file.

**Question I (2 pts)** Write a function that takes a square 2D numpy array along with a number as inputs and returns a 2D diagonal array with the same shape and diagonal element values of the input number. For example, if it takes $\begin{bmatrix} 5 & 15 \\ 100 & 45 \end{bmatrix}$ and value 4, it should return $\begin{bmatrix} 4 & 0 \\ 0 & 4 \end{bmatrix}$.

<!-- grade_id: cell-688a99ee614b4a5a, points: 0, grade: false, solution: true -->
```python
def return_matrix(matrix_in, value_in):
    """
    Inputs:
    - A 2D array with arbitrary size
    - A numeric value

    Output:
    - A same-size diagonal 2D array with diagonal elements of the input numeric value
    """
    # YOUR CODE HERE
    array_size = matrix_in.shape[0] #gets input matrix size

    diagonal_array = np.diag([value_in] * array_size) #creates a diagonal array with matrix N size * value (4)
    return diagonal_array
```

<!-- grade_id: cell-4c27514a2559e9c7, points: 2, grade: true, solution: false -->
```python
# Run this cell to test your code
import numpy as np
array_2d = np.array([[5, 15], [100, 45]])
value_test = 4
assert np.allclose(return_matrix(array_2d, value_test), [[value_test, 0], [0, value_test]])
```

**Question II (13 pts)** For each part, write a function that takes in an arbitrary size 2-dimensional array, and

_Part (1)_ returns a same size array with its elements being the square root of the original array's elements

<!-- grade_id: cell-19f78f41b7982623, points: 0, grade: false, solution: true -->
```python
def square_root_matrix(matrix_in):
    """
    Input:
    A 2D array with arbitrary size

    Output:
    A same-size array with its elements squared values of the original matrix
    """
    # YOUR CODE HERE
    return np.sqrt(matrix_in) #use sqrt function
```

<!-- grade_id: cell-93b27c71e1077bf8, points: 2, grade: true, solution: false -->
```python
# Run this cell to test your code
import numpy as np
array_2d = np.array([[10, 15], [45, 22]])
assert np.allclose(square_root_matrix(array_2d), [[3.16227766, 3.87298335], [6.70820393, 4.69041576]])
```

_Part (2)_ returns the 2D array found by $(A^{T}A)^{-1}$

<!-- grade_id: cell-ff612dbbb8fb5454, points: 0, grade: false, solution: true -->
```python
def invert_func(matrix_in):
    """
    Input:
    A 2D array (matrix)

    Output:
    A^T*A^(-1)
    """
    # YOUR CODE HERE
    A = matrix_in.copy()
    A_p = np.dot(A.T,A)
    Ainv=np.linalg.inv(A_p)
    return Ainv
```

<!-- grade_id: cell-2cbdbe88e28d695b, points: 3, grade: true, solution: false -->
```python
# Run this cell to test your code
import numpy as np
array_2d = np.array([[4, 5], [10, 6]])
assert np.allclose(invert_func(array_2d), [[0.09023669, -0.1183432],[-0.1183432, 0.17159763]])
```

_Part (3)_ returns the sum of all the elements across the columns (axis=0) and standard deviation for each column (axis=1). Pay attention that function should have two outputs and the order for the outputs should be first the sum and then the standard deviation.

<!-- grade_id: cell-9c6b242e981c422f, points: 0, grade: false, solution: true -->
```python
def sum_std(matrix_in):
    """
    Input:
    A 2D array

    Outputs:
    sum of elements across the columns
    Standard deviations down the rows
    """
    # YOUR CODE HERE
    elem_sums = np.sum(matrix_in, axis=0)

    elem_std = np.std(matrix_in, axis=1)

    return elem_sums, elem_std

```

<!-- grade_id: cell-430532aeb67ae384, points: 2, grade: true, solution: false -->
```python
# Run this cell to test your code
import numpy as np
array_2d = np.array([[10, 15], [45, 22]])
assert np.allclose([55, 37], sum_std(array_2d)[0])
assert np.allclose([2.5, 11.5], sum_std(array_2d)[1])
```

_Part (4)_ returns an array with the same size but with sorted elements down the rows (axis=0). For example if it takes
$\begin{bmatrix} 45 & 15 \\ 22 & 10\end{bmatrix}$, it should return $\begin{bmatrix} 22 & 10 \\ 45 & 15\end{bmatrix}$

<!-- grade_id: cell-f1e0764360f286d8, points: 0, grade: false, solution: true -->
```python
def sorted_array(matrix_in):
    """
    Input:
    A 2D array

    Output:
    2D array sorted by the values down the rows
    """
    # YOUR CODE HERE
    sorted_arr = np.sort(matrix_in, axis=0)
    return sorted_arr
```

<!-- grade_id: cell-fc183702638e1242, points: 2, grade: true, solution: false -->
```python
# Run this cell to test your code
import numpy as np
array_2d = np.array([[45, 15], [22, 10]])
assert np.allclose(np.array([[22, 10], [45, 15]]), sorted_array(array_2d))
```

 _Part (5)_ returns a slice that has the elements on rows 2 and 3 and columns 2 to 4 (inclusive). Your function **should first check** if the array has at least 3 rows and 4 columns. If it does not, it should return the message `The array does not have the proper size`. For example, if it takes $\begin{bmatrix} 100 & 150 & 800 & 100 & 130 \\ 500 &  220 &  450 & 3000 & 340 \\ 520 &  550 & 400 & 160 & 150 \end{bmatrix}$, it should return $\begin{bmatrix} 220 & 450 & 3000 \\ 550 & 400 & 160\end{bmatrix}$. If it takes $\begin{bmatrix} 450 & 3000 \\ 400 & 160\end{bmatrix}$, it should return `The array does not have the proper size`.

<font color='blue'>Note:</font> The function should return a string containing `The array does not have the proper size` and NOT print this string. It means you need
```python
return str('The array does not have the proper size')
```


<!-- grade_id: cell-a0a68f584c12f104, points: 0, grade: false, solution: true -->
```python
def slicing_func(matrix_in):
    """
    Input:
    A 2D array

    Output:
    The slice including rows 2 and 3 and columns 2,3, and 4
    """
    # YOUR CODE HERE
    #Checking for an array that is atleast 3x4
    if matrix_in.shape[0] >= 3 and matrix_in.shape[1] >= 4:
      #return rows 2 and 3 and columns 2 to 4 (inclusive)
      elem_slice = matrix_in[1:3, 1:4]
      return elem_slice
    else:
      return 'The array does not have the proper size'
```

<!-- grade_id: cell-78ea5968185aefa2, points: 4, grade: true, solution: false -->
```python
# Run this cell to test your code
# Test 1
import numpy as np
array_2d = np.array([[100, 150, 800, 100, 130], [500, 220, 450, 3000, 340], [520, 550, 400, 160, 150]])
assert np.allclose(np.array([[220,  450, 3000],[550,  400,  160]]), slicing_func(array_2d))
# Test 2
array_2d = np.array([[100, 130], [500, 220], [520, 150]])
assert slicing_func(array_2d)=='The array does not have the proper size'
```

```python

```